<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung in die Programmierung
### Winterersemester 2025/26
Prof. Dr. Stefan Goetze

# Objektorientierte Programmierung (OOP)

## Erstes Beispiel: Erstellen einer Klasse

Wir erstellen im folgenden Code Abschnitt eine einfache Klasse `Hund` mit den Eigenschaften (Attributen) `name` und `rassse`, sowie der Methode `bellen()`.

```python
class Hund:
    def __init__(self, name, rasse):
        self.name = name
        self.rasse = rasse

    def bellen(self):
        return f"{self.name} sagt: Wuff!"
```

**Instanziierung der Klasse:**

Nun können wir eine *Instanz* der Klasse (also eine Variable vom Typ `Hund`) anlegen.

```python
mein_hund = Hund("Bello", "Labrador")

print(mein_hund.bellen())  # Ausgabe: Bello sagt: Wuff!

type(mein_hund)  # Ausgabe: <class '__main__.Hund'>
```

In der Definition der Klasse `Hund` mag "*der Bezeichner*" `self` auffallen, sowie die Funktion `__init__()`.

- `self` stellt eine Referenz auf die aktuelle Instanz der Klasse dar. Wenn Sie ein Objekt einer Klasse erstellen, können Sie mit `self` auf dessen Daten und Funktionen zugreifen.
- `self` muss immer als erstes Argument in den Methoden der Klasse stehen. Dies zeigt an, dass die Methode einer Instanz der Klasse aufgerufen wird.
- Mit `self` können Instanzattribute (Variablen) definieren werden und darauf zugegriffen werden. 
- Die Nutzung des Bezeichners `self` ist (nur) eine Konvention. `self` ist kein reserviertes Wort; Sie könnten es theoretisch auch anders benennen. Es ist jedoch ratsam, bei `self` zu bleiben, um den Code verständlich zu halten.

```python

```

## Die 4 Grundpfeiler der Objektorientierten Programmierung 

1. Kapselung
2. Vererbung
3. Polymorphismus
4. Abstraktion

### Kapselung
Daten (Attribute) und die dazugehörigen Funktionen (Methoden) werden in einer Klasse in einer Einheit gebündelt und der direkten Zugriff auf die Daten von außen wird beschränkt. Dies dient zum
- Schutz der Datenintegrität und zur 
- Reduzierung der Komplexität,
indem Benutzer nur mit definierten Methoden interagieren.

Die folgende Klasse *implementiert ein Bankkonto*. Auf den Kontostand (Variable `__kontostand`) kann nicht direkt zugegriffen werden (nur über die zur Verfügung gestellten Methoden (z.B. `einzahlen()`, `abheben()`,...))

```python
class Bankkonto:
    def __init__(self, kontostand=0):
        self.__kontostand = kontostand  # private Variable

    def einzahlen(self, betrag):
        self.__kontostand += betrag

    def abheben(self, betrag):
        if betrag <= self.__kontostand:
            self.__kontostand -= betrag
        else:
            print("Nicht genügend Guthaben!")

    def anzeigen(self):
        return self.__kontostand
```

```python
mein_konto = Bankkonto()
print(mein_konto.anzeigen())  # Ausgabe: 0
mein_konto.einzahlen(100)
print(mein_konto.anzeigen())  # Ausgabe: 100
mein_konto.abheben(50)
print(mein_konto.anzeigen())  # Ausgabe: 50
mein_konto.abheben(100)  # Ausgabe: Nicht genügend Guthaben!
```

Auf die Variable `__kontostand` der Klasse `Bankkonto` kann nicht von außerhalb der Klasse zugegriffen werden, da diese *privat* ist (dazu später noch mehr).

```python
print(mein_konto.__kontostand)  # Fehler: AttributeError, da __kontostand privat ist
```

#### Verwendung von Unterstrichen in Python

Die doppelten Unterstriche in Python, wie z.B. bei `__init__`, haben spezielle Bedeutungen und dienen hauptsächlich der Namensverknüpfung (Name Mangling) und der Vermeidung von Namenskonflikten. Hier sind die wichtigsten Punkte:

1. Spezielle Methoden (Dunder-Methoden)
   - Die doppelten Unterstriche markieren spezielle Methoden, die oft als „Dunder-Methoden“ (Double Underscore Methods) bezeichnet werden. Beispielhafte Methoden sind:
     - `__init__`: Der Konstruktor, der beim Erstellen eines Objekts aufgerufen wird.
     - `__str__`: Definiert, wie das Objekt als String dargestellt wird.
     - `__repr__`: Gibt eine offizielle String-Repräsentation des Objekts zurück.
   - Diese Methoden haben spezifische Aufgaben und können von Python intern aufgerufen werden. 
   - Die doppelten Unterstriche sind nicht nur eine Konvention, sondern ein wichtiges Merkmal der Python-Programmierung, das die Struktur und Sicherheit von Code verbessert. Sie führen zu einer klareren Trennung zwischen öffentlichen und privaten Attributen und Methoden.

2. Namensverknüpfung (Name Mangling)
   - Bei Attributen, die mit doppelten Unterstrichen definiert sind (z.B. __attribute), wandelt Python den Namen intern um, um Namenskonflikte in abgeleiteten Klassen zu vermeiden. Das bedeutet, dass:
    - __attribute in der Klasse MyClass intern als _MyClass__attribute gespeichert wird.
Dies macht es schwierig, direkt auf diese Attribute von außen zuzugreifen, fördert aber den sicheren Zugriff auf die Attribute innerhalb der Klasse.

3. Kapselung von Attributen
   - Durch die Verwendung von doppelten Unterstrichen wird eine Form von Kapselung erreicht. Das ermöglicht es, interne Implementierungsdetails eines Objekts zu verbergen.

Beispiel
Hier ist ein einfaches Beispiel, das die Verwendung von `__init__` und doppelten Unterstrichen verdeutlicht:

```python
class BeispielKlasse:
    def __init__(self, wert):
        self.__wert = wert  # private Attribut mit Namensverknüpfung
        self._wert = wert   # geschütztes Attribut
        self.wert = wert    # öffentliches Attribut

    def zeige_wert(self):
        return self.__wert

objekt = BeispielKlasse(10)
print(objekt.zeige_wert())  # Ausgabe: 10
```

### Platzierung von Unterstrichen

- Am Anfang eines Bezeichners:

  - `self._attribut`: wird als **protected** betrachtet, d.h. es soll nur innerhalb der Klasse und von abgeleiteten Klassen verwendet werden.
  - `self.__attribut`: wird durch Namensverknüpfung geschützt und ist **privat**“** für die Klasse.

- Am Ende eines Bezeichners:

  - `name_` oder ``name__: Hierbei handelt es sich normalerweise um die Konvention, um Reservierungskonflikte mit Schlüsselwörtern oder um spezielle Funktionen zu vermeiden. Der Unterstrich am Ende wird verwendet, um Konflikte mit eingebaute Namen zu vermeiden.


```python
print(objekt._wert)  # Dies sollte man nicht tun, aber es funktioniert
print(objekt.wert)   # Öffentliches Attribut, kein Problem
```

```python
print(objekt.__wert)  # Dies führt zu einem Fehler
```

### Vererbung

```python
class Mops(Hund):
    def bellen(self):
        return f"{self.name} sagt: Wuff! (aber anders)"
```

### Polymorphismum

```python
def hund_bellen(hund):
    print(hund.bellen())

hunde = [mein_hund, Mops("Fido", "Mops")]
for h in hunde:
    hund_bellen(h)
```

**Übung 1:**<br>
Schreiben Sie eine Klasse `Standort`. Die Klasse soll die Daten *Postleitzahl* (`zipcode`), *Längengrad* (`longitude`) und Breitengrad (`latitude`) speichern. <br> Zudem soll die Klasse eine Funktion `getCoordinates` enthalten, welche die `latitude` und die `longitude` als Tupel zurückgibt.<br>
Der Konstruktor von Standort soll die Parameter `zipcode`, `longitude` und `latitude` entgegennehmen und in den entsprechenden Eigenschaften ablegen.

<!-- grade_id: cell-e6d754d6e230ab48, points: 0, grade: false, solution: true -->
```python
class Standort:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ce79c3192186d96e, points: 1, grade: true, solution: false -->
```python
"""test class standort"""
from nose.tools import assert_equal

s = Standort("1111", 1.0, 1.0)
assert_equal((1.0, 1.0), s.getLocation(), "Da stimmt etwas mit der Zuweisung nicht!")

```

Der offene Webservice zippopotam.us stellt Informationen zu Postleitzahlen in einer vielzahl von Ländern im JSON-Format zur Verfügung.

Die Route um Daten abzurufen ist wie folgt aufgebaut:
*https://api.zippopotam.us/[country]/[zipcode]"*


**Übung 2:** <br>

Bauen Sie das Programm aus der <u>Übung 2 von Arbeitsblatt `webservices.ipynb`</u> so um, dass ein <u>Objekt vom Typ Standort</u> mit den erhaltenen Daten zurückgegeben wird.

*Tipp: Der Country-Code für Deutschland lautet: de*<br>
*Tipp: Interpretieren Sie die JSON Daten mit json.load([daten])*

<!-- grade_id: cell-e7aab6c95906bfb5, points: 0, grade: false, solution: true -->
```python
import json
import urllib

def fetch(country, zipcode):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0a5d4c0199321ce1, points: 1, grade: true, solution: false -->
```python
"""Tests fetch"""
from nose.tools import assert_equal

s = fetch("de", "58644")

assert isinstance(s,Standort), "Die Funktion 'fetch' soll ein Objekt vom Typ 'Standort' zurückliefern"
assert_equal(('7.6968', '51.3246'), s.getLocation(), "Es wurden die falschen Daten extrahiert.")
```