<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung in die Programmierung
### Winterersemester 2025/26
Prof. Dr. Stefan Goetze

# Datei-Handling

Bevor wir mit Dateien arbeiten können laden wir und erst einmal ein paar Dateien aus dem Internet herunter. Hierzu benutzen wir die Bibliotheken [`urllib`](https://docs.python.org/3/library/urllib.html) für den Zugriff auf's Internet (Download) und [`os`](https://docs.python.org/3/library/os.html) um das Betriebssystem anzuweisen, zu überprüfen, ob die zu speichernde Datei ggf. schon existiert.

```python
from urllib.request import urlopen            # to download a file
from urllib.error import URLError, HTTPError  # to handle errors
import os.path                                # to check if file exists 

def download_file(url, fname):
    '''Downloads a file from the given URL if it does not already exist locally.
    Args:
        url (str): The URL of the file to download.
        fname (str): The local filename to save the downloaded file as.
    '''
    # download the file if it does not exist
    if not os.path.exists(fname):
        try:
            with urlopen(url) as response:
                content = response.read().decode('utf-8')
                with open(fname, "w") as file:
                    file.write(content)
            print(f"Downloaded file {fname}.")
        except HTTPError as e:
            print(f"HTTP error occurred: {e.code} - {e.reason}")
        except URLError as e:
            print(f"URL error occurred: {e.reason}")
        except Exception as e:
            print(f"An unexpected error occurred: {str(e)}")
    else:
        print(f"File '{fname}' already exists, skipping download.")

download_file("https://raw.githubusercontent.com/fh-swf-hgi/py4e/master/code3/mbox-short.txt", "mbox-short.txt")
download_file("https://raw.githubusercontent.com/fhswf/EinfProg/refs/heads/main/u5/example.csv", "example.csv")
download_file("https://raw.githubusercontent.com/fhswf/EinfProg/refs/heads/main/u5/example.json", "example.json")
download_file("https://raw.githubusercontent.com/fhswf/EinfProg/refs/heads/main/u5/example.xml", "example.xml")
```

## Dateien öffnen

Im obigen Beispiel zum herunterladen und speichern der Datei auf unserer Festplatte haben wir bereits gesehen, dass der Befehl `read()` zum Lesen von Dateiinhalten und der Befehl `write()` zum Schreiben von Dateiinhalten verwendet werden kann. Wir wollen uns nun zuerst das *Öffnen* und *Lesen von Dateien* genauer ansehen.

Wir öffnen eine Datei mit dem Befehl [`open()`](/https://docs.python.org/3/library/functions.html#open), der uns ein [Datei-Objekt](https://docs.python.org/3/glossary.html#term-file-object) zurück gibt (hier die Variable `file`). Neben dem Dateinamen `fname` geben wir noch den Modusparameter `r` für *Lesen* (engl: read) an.

Der *Modus* (`mode`), in dem die Datei geöffnet wird, beeinflusst, wie file verwendet werden kann. Folgende Modi können (auch in Kombiunation) genutzt werden
- `'r'`: Lesen (*default*)
  - Öffnet eine Datei zum Lesen. Die Datei muss existieren; andernfalls wird ein Fehler ausgegeben.
- `'w'`: Schreiben 
  - Öffnet eine Datei zum Schreiben. Wenn die Datei bereits existiert, wird sie überschrieben (*truncated*). Andernfalls wird eine neue Datei erstellt.
- `'a'`: Anhängen
  - Öffnet eine Datei zum Anhängen. Neue Daten werden am Ende der Datei hinzugefügt. Wenn die Datei nicht existiert, wird sie erstellt.
- `'x'`: Exklusive Erstellung
  - Öffnet eine Datei zum Schreiben, schlägt jedoch fehl, wenn die Datei bereits existiert.
- `'+'`: Öffnen zum Lesen und Schreiben
- `'b'`: Binärmodus (binary)
  - Öffnet die Datei im Binärmodus, für nicht-textuellen Daten (z.B. Bildern oder Audiodateien) gearbeitet wird. Kann mit anderen Modi kombiniert werden (z.B. `'rb'` oder `'wb'`).
- `'t'`: Textmodus (*default*)

Wird kein Modus angegeben ist `mode='rt'`, also Lesen im Textmodus (man könnte den Parameter also im nachfolgenden Codeabschnitt auch weg lassen).

Beim Öffnen einer Textdatei kann ein spezifisches Encoding angegeben werden, z.B. [UTF-8](https://de.wikipedia.org/wiki/UTF-8) (`encoding='utf-8'`). Dies ist wichtig wenn nicht-ASCII Zeichen im Text verwendet werden (sollen).

```python
fname = "mbox-short.txt"

# Datei öffnen 
file = open(fname, 'r', encoding="utf-8") 
```

## Dateien lesen

Der Datei-Handler `file` im vorherigen Codeabschnitt ist ein objektherausgegebenes Dateiobjekt (Typ: `TextIOWrapper` oder `BufferedReader`/`BufferedWriter`, je nach Modus).

Nachfolgend ein paar möglicherweise interessante [Methoden des Datei-Handlers](https://docs.python.org/3/tutorial/inputoutput.html#methods-of-file-objects):
- `read(size)`: Liest die Datei und gibt den Inhalt zurück. Die Option size gibt an, wie viele Bytes oder Zeichen gelesen werden sollen.
- `readline()`: Liest eine Zeile aus der Datei und gibt sie zurück.
- `readlines()`: Liest alle Zeilen der Datei und gibt sie als Liste zurück.
- `write(string)`: Schreibt die angegebene Zeichenkette in die Datei.
- `writelines(list)`: Schreibt eine Liste von Zeichenfolgen in die Datei.
- `seek(offset, whence)`: Setzt den Lese-/Schreibzeiger in der Datei.
- `tell()`: Gibt den aktuellen Wert des Lese-/Schreibzeigers zurück.

```python
# Dateiinhalt lesen
content = file.read() 

# Dateiinhalt ausgeben
print(content)
```

```python
# Option 1: Lesezeiger zurücksetzen und erneut lesen
file = open(fname, 'r', encoding="utf-8") 

file.seek(0)                    # auf Anfang zurücksetzen
content2 = file.readlines(70)   # jetzt werden wieder Zeilen geliefert
print(content2)           # Ausgabe der Zeilen 5 bis 59
```

```python
from itertools import islice

# Demonstration: why readlines(hint) appears insensitive to the hint

fname = "mbox-short.txt"

# readlines(hint) reads whole lines until approx hint bytes accumulated
with open(fname, 'r', encoding='utf-8') as f:
    chunk_lines_20 = f.readlines(20)
print("readlines(20) ->", len(chunk_lines_20), "lines, first line len:", len(chunk_lines_20[0]) if chunk_lines_20 else 0)
print(chunk_lines_20,"\n")

with open(fname, 'r', encoding='utf-8') as f:
    chunk_lines_55 = f.readlines(55)
print("readlines(55) ->", len(chunk_lines_55), "lines, first line len:", len(chunk_lines_55[0]) if chunk_lines_55 else 0)
print(chunk_lines_55,"\n")

with open(fname, 'r', encoding='utf-8') as f:
    chunk_lines_57 = f.readlines(57)
print("readlines(57) ->", len(chunk_lines_57), "lines, first line len:", len(chunk_lines_57[0]) if chunk_lines_57 else 0)
print(chunk_lines_57,"\n")

with open(fname, 'r', encoding='utf-8') as f:
    chunk_lines_70 = f.readlines(70)
print("readlines(70) ->", len(chunk_lines_70), "lines, first line len:", len(chunk_lines_70[0]) if chunk_lines_70 else 0)
print(chunk_lines_70,"\n")

# If you want exact bytes, use read(n)
with open(fname, 'r', encoding='utf-8') as f:
    exact_bytes = f.read(20)
print("read(20) ->", repr(exact_bytes), "length:", len(exact_bytes))
print(exact_bytes,"\n")

# If you want an exact number of lines, use islice
with open(fname, 'r', encoding='utf-8') as f:
    first_3_lines = list(islice(f, 3))
print("first 3 lines (islice) ->", len(first_3_lines))
```

```python
# Option 2: Datei neu öffnen (sicherer, bevorzugt)
with open(fname, 'r', encoding='utf-8') as f:
    content2 = f.readlines(10)
    print(content2)
```

## Datei schließen

Dateien sollten immer nach ihrer Verwendung geschlossen werden, um Ressourcen freizugeben und mögliche Datenverluste zu vermeiden. Dabei wird oft die Methode `close()` verwendet.

```python
# Datei schließen
file.close()
```

Alternativ ist es empfehlenswert, den `with`-Kontextmanager zu verwenden, der die Datei automatisch schließt, auch wenn ein Fehler auftritt. Hier ist ein Beispiel, wie man die Datei richtig schließt:

```python
# Datei mit 'with' öffnen
with open(fname, 'r') as file:
    content = file.read()
# Datei wird automatisch geschlossen

#print(content)
```

## Fehlerbehandlung bei der Arbeit mit Dateien

Wie schon im letzten Übungs-Notebook besprochen, ist es wichtig, beim Arbeiten mit file mögliche Fehler zu berücksichtigen, z.B. mit `try` - `except`, um Fehler wie „Datei nicht gefunden“ zu handhaben.

Hier ist ein einfaches Beispiel, das die bisher besprochenen Konzepte zusammenführt:

```python
try:
    with open(fname, 'r', encoding='utf-8') as file:
        content = file.read()
        print(content)
        print("Aktuelle Position:", file.tell())
except FileNotFoundError:
    print("Die Datei wurde nicht gefunden.")
except IOError:
    print("Ein Fehler ist beim Zugriff auf die Datei aufgetreten.")
```

## Weitere Text-basierte Dateien

### CSV Dateien (Comma-Separated Values)

[CSV-Dateien (Comma-Separated Values)](https://de.wikipedia.org/wiki/CSV_(Dateiformat)) sind eines der am häufigsten verwendeten Dateiformate zur Speicherung und zum Austausch von Daten. CSV-Dateien sind im Textformat und leicht verständlich. Sie können mit jedem Texteditor geöffnet und bearbeitet werden.

Viele Anwendungen wie Microsoft Excel und Google Sheets unterstützen CSV-Dateien. Dies macht sie zu einem Standardformat für den Datenaustausch.

In vielen Datenanalysesystemen und -bibliotheken (z.B. [`Pandas`](https://pandas.pydata.org/) in Python) gibt es integrierte Funktionen zum einfachen Importieren und Verarbeiten von CSV-Daten.

CSV-Dateien sind im Vergleich zu komplexeren Formaten wie JSON oder XML speicherplatzsparend und ermöglichen eine schnelle Verarbeitung großer Datenmengen.

Wir verwendet zum Lesen einer CVS Datei die Bibliothek [`csv`](https://docs.python.org/3/library/csv.html).

Der Inhalt der CSV-Datei `'example.csv'` ist wie folgt:

```CSV
Name,Geschlecht,Stadt
Alice,w,Berlin
Bob,m,München
Stefan,m,Iserlohn
```

```python
import csv

# Öffnen der CSV-Datei im Lesemodus
with open('example.csv', 'r', encoding='utf-8') as csvfile:
    reader = csv.reader(csvfile)

    print("Inhalt der CSV-Datei:")
    print("-----------------------")
    
    # Überspringen der Kopfzeile 
    next(reader)
    
    # Durch die Zeilen in der CSV-Datei iterieren
    for row in reader:
        name = row[0]
        geschlecht = row[1]
        stadt = row[2]

        print(f'Name: {name:6}, Geschlecht: {geschlecht}, Stadt: {stadt}')
```

Da CSV Dateien eigentlich Textdateien sind und wir ja schon Experten im Programmieren sind (zumindest auf dem Weg dahin), können wir die CSV Datei natürlich auch ohne die `csv` Bibliothek einlesen. Hierzu benutzen wir die Befehle [`readline()`](https://docs.python.org/3/tutorial/inputoutput.html#methods-of-file-objects) um jeweils eine Zeile der Textdatei einzulesen, [strip()](https://docs.python.org/3/library/stdtypes.html#str.strip) um Leerzeichen und Zeilenumbrüche zu entfernen und [split(',')] um eine Zeile in Ihre Werte (bei jedem `,`) aufzuteilen.

```python
# Öffnen der CSV-Datei im Lesemodus
with open('example.csv', 'r', encoding='utf-8') as file:
    # Lesen der Datei Zeile für Zeile
    header = file.readline().strip()  # Kopfzeile lesen
    print("Kopfzeile:", header)
    
    # Iteriere über die restlichen Zeilen
    for line in file:
        # Entfernen von Leerzeichen und Zeilenumbrüchen
        line = line.strip()
        
        # Aufteilen der Zeile in Komponenten
        zeile = line.split(',')
        
        # Ausgabe der Werte
        name = zeile[0]
        geschlecht = zeile[1]
        stadt = zeile[2]
        print(f'Name: {name:6}, Geschlecht: {geschlecht}, Stadt: {stadt}')
```

### JSON Dateien

[JSON (JavaScript Object Notation)](https://de.wikipedia.org/wiki/JSON) ist ein leistungsstarkes und vielseitiges Format zum Speichern und Übertragen von Daten, das in der Softwareentwicklung, insbesondere bei Webanwendungen, einen hohen Stellenwert hat. Die Vorteile von JSON machen es zu einer der am häufigsten verwendeten Datenformatierungen in der heutigen datengetriebenen Welt.


```JSON
[
    {
        "name": "Alice",
        "alter": 30,
        "stadt": "Berlin"
    },
    {
        "name": "Bob",
        "alter": 25,
        "stadt": "München"
    },
    {
        "name": "Stefan",
        "alter": 15,
        "stadt": "Iserlohn"
    }
]
```

Wie wir sehen, ist CSV kompakter als JSON, aber JSON ermöglicht die Verwendung von Schlüsseln, die dazu beitragen, Daten leicht zu identifizieren und zu extrahieren. Wie zu erwarten gibt es in Python eine mächtige Bibliothek zur Verarbeitung von JSON Dateien ([`json`](https://docs.python.org/3/library/json.html)).

```python
import json

# Öffnen der JSON-Datei im Lesemodus
with open('example.json', 'r', encoding='utf-8') as jsonfile:
    # Laden der Daten aus der JSON-Datei
    daten = json.load(jsonfile)
    
    # Durch die Liste der Daten iterieren und ausgeben
    for eintrag in daten:
        name = eintrag['name']
        alter = eintrag['alter']
        stadt = eintrag['stadt']
        print(f'Name: {name:6}, Alter: {alter}, Stadt: {stadt}')
```

### XML Dateien

Falls die Daten als [Extensible Markup Language (XML)-Datei](https://de.wikipedia.org/wiki/Extensible_Markup_Language) vorliegen ist dies natürlich auch kein Problem, auch hierfür bietet Python eine Bibliothek.

```python
import xml.etree.ElementTree as ET

# Einlesen der XML-Datei
tree = ET.parse('example.xml')
root = tree.getroot()

# Durch die Einträge iterieren und Informationen ausgeben
for eintrag in root.findall('eintrag'):
    name = eintrag.find('name').text
    alter = eintrag.find('alter').text
    stadt = eintrag.find('stadt').text
    print(f'Name: {name:6}, Alter: {alter}, Stadt: {stadt}')
```

## Schreiben von Dateien

Da die vorhin behandelte CSV Datei `'example.csv'` 

```CSV
Name,Geschlecht,Stadt
Alice,w,Berlin
Bob,m,München
Stefan,m,Iserlohn
```

ja nichts anderes als eine Textdatei können wir hier sehr einfach (ohne weitere Bibliotheken) einen Datensatz (eine Zeile hinzufügen):

```python
# Neue Daten, die hinzugefügt werden sollen
neuer_name = "Charlie"
neues_geschlecht = "w"
neue_stadt = "Altena"

# Öffnen der CSV-Datei im Anhängmodus
with open('example.csv', 'a', encoding='utf-8') as file:
    # Erstellen der neuen Zeile als Zeichenkette
    neue_zeile = f"{neuer_name},{neues_geschlecht},{neue_stadt}\n"
    
    # Schreiben der neuen Zeile in die Datei
    file.write(neue_zeile)
```

Lassen Sie uns nachsehen, ob das Anhängen des Datensatzes geklappt hat:

```python
# Öffnen der CSV-Datei im Lesemodus
with open('example.csv', 'r', encoding='utf-8') as csvfile:
    reader = csv.reader(csvfile)

    print("Inhalt der CSV-Datei:")
    print("-----------------------")
    
    # Überspringen der Kopfzeile 
    next(reader)
    
    # Durch die Zeilen in der CSV-Datei iterieren
    for row in reader:
        name = row[0]
        geschlecht = row[1]
        stadt = row[2]

        print(f'Name: {name:6}, Geschlecht: {geschlecht}, Stadt: {stadt}')
```

## Suchen in Dateien

Nutzung des `in` Keywords (sowohl in der `for`-Schleife als auch in der `if`-Abfrage).

```python
def suchen_in_datei(dateipfad, suchbegriff):
    try:
        # Öffnen der Datei im Lesemodus
        with open(dateipfad, 'r', encoding='utf-8') as file:
            zeilen = file.readlines()
            gefunden = False
            
            # Durch die Zeilen iterieren und nach dem Suchbegriff suchen
            for index, zeile in enumerate(zeilen):
                if suchbegriff.lower() in zeile.lower():
                    print(f"Gefunden in Zeile {index + 1}: {zeile.strip()}")
                    gefunden = True

            if not gefunden:
                print(f"Der Begriff '{suchbegriff}' wurde in der Datei nicht gefunden.")
                
    except FileNotFoundError:
        print(f"Die Datei '{dateipfad}' wurde nicht gefunden.")
    except Exception as e:
        print(f"Ein Fehler ist aufgetreten: {str(e)}")

# Aufruf der Suchfunktion
#suchen_in_datei("mbox-short.txt", "From")
suchen_in_datei("mbox-short.txt", " m04")
```

Anstelle der Verwendung des Operators `in` kann auch die Funktion [`find()`](https://docs.python.org/3/library/stdtypes.html#str.find) zur Suche in Texten verwendet werden. Im nachfolgenden Code-Beispiel wird in Zeile $9$ `zeile.lower().find(suchbegriff.lower())` verwendet. Wenn der Suchbegriff gefunden wird, gibt `find()` den Index des ersten Auftretens zurück; andernfalls gibt es $-1$ zurück.

```python
def suchen_in_datei(dateipfad, suchbegriff):
    try:
        # Öffnen der Datei im Lesemodus
        with open(dateipfad, 'r', encoding='utf-8') as file:
            gefunden = False
            
            # Durch die Zeilen iterieren und nach dem Suchbegriff suchen
            for index, zeile in enumerate(file):
                if zeile.lower().find(suchbegriff.lower()) != -1:
                    print(f"Gefunden in Zeile {index + 1}: {zeile.strip()}")
                    gefunden = True

            if not gefunden:
                print(f"Der Begriff '{suchbegriff}' wurde in der Datei nicht gefunden.")
                
    except FileNotFoundError:
        print(f"Die Datei '{dateipfad}' wurde nicht gefunden.")
    except Exception as e:
        print(f"Ein Fehler ist aufgetreten: {str(e)}")

# Aufruf der Suchfunktion
suchen_in_datei("mbox-short.txt", "From")
```

## Übungen

**Übung 1: Schreiben Sie eine Funktion `lese_datei`, die eine Datei einliest und den Inhalt der Datei in Großbuchstaben (Zeile für Zeile) ausgibt.**

<!-- grade_id: cell-c1e55762bc63dc93, points: 0, grade: false, solution: true -->
```python
def lese_datei(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
with open('myfile.txt', "w") as myfile:
    myfile.write("Ein neuer\nText.\nEnde")
lese_datei("myfile.txt")
```

<!-- grade_id: cell-0aa5a4884c58af6f, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal('EIN NEUER\n\nTEXT.\n\nENDE\n', str(cap))
```

**Übung 2: Schreiben Sie eine Funktion `berechne_durchschnitt`, die eine Datei einliest und in der Datei nach Zeilen der Form 
`X-DSPAM-Confidence: 0.8475` sucht. Die Zahl kann sich dabei ändern, die anderen Zeichen nicht. Insbesondere muss die Zeile immer mit `X-DSPAM...` beginnen.<br>Extrahieren Sie die Fließkommerzahl indem Sie die Zeile auftrennen. Berechnen Sie die Summe der Fließkommerzahlen und geben Sie den durchschnittlichen Wert der Fließkommerzahl zurück.**

<!-- grade_id: cell-8e47dfd8aabc8a20, points: 0, grade: false, solution: true -->
```python
def berechne_durchschnitt(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ead248e87efeb5f, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal

# Durchschnittswert prüfen
assert_almost_equal(0.7507185185185187, berechne_durchschnitt(fname))
```